---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: dc909ab842
xp: 50
```

`@projector_key`
39c10b60df858b5cb1b3153613cac970

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 774d92a419
xp: 50
```

`@projector_key`
ff0a62232a78b17e63bdfdee751b599e
